# Artificial Intelligence: Concepts, Challenges, and Opportunities (2020), exercises


## General instructions for all exercises

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Follow the instructions and fill in your solution under the line marked by tag

> YOUR CODE HERE
  
Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached the code cell which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

The notebook may also contain manually graded answers. Write your manualle graded answer under the line marked by tag:

> YOUR ANSWER HERE

Manually graded tasks may be text, pseudocode, or mathematical formulas. You can write formulas with $\LaTeX$-syntax by enclosing the formula with dollar signs (`$`), for example `$f(x)=2 \pi / \alpha$`, will produce $f(x)=2 \pi / \alpha$

When you have passed the tests in the notebook, and you are ready to submit your solutions, download the whole notebook, using menu `File -> Download as -> Notebook (.ipynb)`. Save the file in your hard disk, and submit it in [Moodle](https://moodle.uwasa.fi) under the corresponding excercise.

Your solution should be an executable Python code. Use the code already existing as an example of Python programing and read more from the numerous Python programming material from the Internet if necessary. 


```python
NAME = ""
Student_number = ""
```

---

# Solving 8-puzzle with A*

The possible moves in the 8-puzzle game forms a graph. Obviously, if the tree is exanded from all possible moves, the goal state will be included in the tree. The goal can be find using graph search algorithms, such as A*.

![EightPuzzleGraph](8puzzle-tree.svg)

This 8-puzzle is still quite small, in general the number of possible moves will grow exponentially, and the resulting graph may become very large. Most of the branches are, however, bad moves and not necessary to include in the search graph. Therefore good strategy is to use a heuristic function to assess the quality of the possibile branches with some heuristic method before expanding the branch further. This kind of expanding-on-the-go is not directly supported by networkx library, and therefore this is better to solve using tailor made implementation of the search algorithm. 

In this example an implemenation included in the course book is used:
   
> Python code for the book *[Artificial Intelligence: A Modern Approach](http://aima.cs.berkeley.edu).* You can use this in conjunction with a course on AI, or for study on your own. We're looking for [solid contributors](https://github.com/aimacode/aima-python/blob/master/CONTRIBUTING.md) to help.

Import some libraries and define a simple function for displaying the puzzle.

```python
from search import EightPuzzle, recursive_best_first_search, astar_search
import numpy as np
import ipywidgets as widgets
from IPython.display import display
```

```python
def displayPuzzle(state):
    Bstyle=widgets.Layout(width='30px',)
    items = [widgets.Button(description=str(i), layout=Bstyle) for i in state]
    display(widgets.GridBox(items, layout=widgets.Layout(grid_template_columns="repeat(3, 35px)")))
```

Set up an initial state and goal state, and display them.

```python
print("Intial state:")
initial_state=(2, 4, 3, 1, 5, 6, 7, 8, 0)
displayPuzzle(initial_state)

print("\nGoal state:", end='')
goal_state=(1, 2, 3, 4, 5, 6, 7, 8, 0)
displayPuzzle(goal_state)

puzzle = EightPuzzle(initial_state)
```

## Task 1: Define linear heuristic function

Implement a simple linear heuristic function, which simply calculates the number of tiles in wrong position in the current state. The current state and the goal state are given as arguments to the function. Both are given as simple vector.

The search algorithm uses this heuristic function when eavaluating each node. Based on the result of the heuristic function, the search is directed towards optimal paths.

<!-- grade_id: cell-6a16170ec273b999, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3acd37f4df6cde58, points: 3, grade: true, solution: false -->
```python
distance=linear(initial_state, goal_state)
def heuristic(node):
    return linear(node.state, goal_state)

path=astar_search(puzzle, heuristic).solution()
moves_needed=len(path)
print("Distance=", distance, ", moves=", moves_needed)
print("Path=", path)

assert(linear(initial_state, goal_state)==3)
assert(moves_needed == 8)
```

## Task 2: Define heuristic function using Manhattan distance

**You can choose to answer either Task 2 or Task 4. If you do not want to answer to this question, write your answer on Task 4!**

Implement heuristic function, which return is the sum of manhattan distances of each tile from the correct position. For example in the initial state the distance from the goal state would be:

|Tile | horizontal distance | vertical distance | total distance |
|-----|------|-----|------|
|      1  | 0 |                   1    |              1 |
|      2  | 1 |                   0    |              1 |
|      3  | 0 |                   0    |              0 |
|      4  | 1 |                   1    |              2 |
|      5  | 0 |                   0    |              0 |
|      6  | 0 |                   0    |              0 |
|      7  | 0 |                   0    |              0 |
|      8  | 0 |                   0    |              0 |




The rest of the tiles are in correct places. The total Manhattan distance is in this case 1+1+2 = 4.


<!-- grade_id: cell-486509a22a3a1023, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-09a1e4830c0d8e61, points: 4, grade: true, solution: false -->
```python
distance=manhattan(initial_state, goal_state)
def heuristic(node):
    return manhattan(node.state, goal_state)

path=astar_search(puzzle, heuristic).solution()
moves_needed=len(path)
print("Distance=", distance, ", moves=", moves_needed)
print("Path=", path)

assert(manhattan(initial_state, goal_state)==4)
assert(moves_needed == 8)
```

```python

```

## Task 3. Running time

Whatever heuristic function is used, the algorithm seems to find the same path. That is because it uses recursive search, to search over several possible paths, and chooses the best one. But the heuristic function affects how fast the solutions are found. 

The `%timeit` operator in Jupyter notebooks can be used in finding out how long does it take to accomplish an algorithm. It usullay runs the algorith many times, like 100 or 1000, and estimates the average running time. Take a look at the running times reported below. 

```python
# Estimate the time of finding the solution without heuristic function

def heuristic(node):
    return 0
%timeit astar_search(puzzle, heuristic).solution()

## 18.4 ms ± 11.9 µs
```

```python
# Estimate the time of finding the solution with linear heuristic function

def heuristic(node):
    return linear(node.state, goal_state)
%timeit astar_search(puzzle, heuristic).solution()

## 535 µs ± 1.27 µs
```

```python
# Estimate the time of finding the solution with manhattan heuristic function
def heuristic(node):
    return manhattan(node.state, goal_state)
%timeit astar_search(puzzle, heuristic).solution()

## 944 µs ± 609 ns
```

Answer the following questions
 1. How and why is the heuristic function affecting to the running times?
 1. How can you explain the differences with different heuristic functions?
 1. How could you optimize the heuristic function for a given task?

<!-- grade_id: cell-a1b7256fb7efc63c, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

## Task 4

**If you do not want to answer to task 2, you can answer to this task instead. But choose either one!**

#### Graph search algorithms: 

Select three different graph searching algorithms and explain how they work, and when they are appropriate to use, and when they are not probably a good choise?

<!-- grade_id: cell-b75671732b435a4d, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE