# Artificial Intelligence: Concepts, Challenges, and Opportunities (2021), exercises


## Material for Quiz T10

Start from the beginning of the notebook and execute each code cell by first activating it and then pressing `Shift-Enter` keys at the same time. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached a cell which contains a tag:

> YOUR CODE HERE

Follow the instructions and fill in your solution under the line marked by tag.

Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. Then execute in a similar way a cell below the input cell, which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

When you are finished working with the notebook, download the whole notebook by using menu `File -> Download`. Save the file in your hard disk in case you need it later, because everything in the binder will be destroyed when the server expires.

When you are finished working with the notebook, you should be ready to answer to the questions in the Quizz.



---

# Probabilisic reasoning and machine learning


```python
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()

```

## Task 1, Probabilistic reasoning.

Hundred apples and hundred onions have been accidentally mixed in the grocery store. A shopkeeper developed a robot which can separate onions from apples by simply measuring their weights, since onions are usually lighter than apples. But not always. The following figure shows the weight distribution of apples and onions. 

```python
from numpy import mean , linspace
import pomegranate as pg

onions=pg.NormalDistribution(110,25)
apples=pg.NormalDistribution(170,28)

# Combined, all students
combined = pg.GeneralMixtureModel([apples, onions])

i=linspace(0,250)
plt.plot(i,onions.probability(i), label='Onions')
plt.plot(i,apples.probability(i), label='Apples')
plt.plot(i,combined.probability(i), label='Combined')
plt.xlabel('Weight')
plt.title('Weight of Onions and Apples')
plt.legend()
```

The separation of apples from onions can be implemented using Bayesian classifier. Use the Bayes rule with the help of the pomegranate tools to find out what is the probability that an item is an onion if it weights 125 g ?

- You can make exact calculations using the Pomegranate distribution objects. In this case, make sure that the final calculation is on the last line of the cell, so that it will be the output value of the cell.
- You can make some tests of your answer by examining if the assert on the next cell is passed

<!-- grade_id: cell-b661074d6b5a17e9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4006ae7503c7e866, points: 2, grade: true, solution: false -->
```python
ans = _
print("Your answer is %f" % (ans))
if (ans<=0) or (ans>=1):
    print("Warning! The probability must be between 0 and 1!" )
assert((ans>0) and (ans<1)), "%d is not in acceptable range" % (ans)

```

<!-- grade_id: cell-9011bedc7c97e1ea, points: 0, grade: false, solution: true -->
```python

```