# Artificial Intelligence: Concepts, Challenges, and Opportunities (2020), exercises


## General instructions for all exercises

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Follow the instructions and fill in your solution under the line marked by tag

> YOUR CODE HERE
  
Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached the code cell which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

The notebook may also contain manually graded answers. Write your manualle graded answer under the line marked by tag:

> YOUR ANSWER HERE

Manually graded tasks may be text, pseudocode, or mathematical formulas. You can write formulas with $\LaTeX$-syntax by enclosing the formula with dollar signs (`$`), for example `$f(x)=2 \pi / \alpha$`, will produce $f(x)=2 \pi / \alpha$

When you have passed the tests in the notebook, and you are ready to submit your solutions, download the whole notebook, using menu `File -> Download as -> Notebook (.ipynb)`. Save the file in your hard disk, and submit it in [Moodle](https://moodle.uwasa.fi) under the corresponding excercise.

Your solution should be an executable Python code. Use the code already existing as an example of Python programing and read more from the numerous Python programming material from the Internet if necessary. 


```python
NAME = ""
Student_number = ""
```

---

## Excercise 1: Intelligent agents and utility

## Simple reflex agent

Below is a simple example of autonomous car, which drives in the straight road. The car has simple sensors by which it can observe if it is just about to cross the solig border line on the right hand side or the dashed line on the left hand side of the lane. The steering wheel is the actuator of the car, which can be used to steer the car to the left or to the right. 

Design an intelligent agent, which can steer the autonomous car trough the road staying in the right lane all the time using $\tt{percepts} \in \{\tt{ left\; |\; middle\; |\; right }\}$ and using the actuator commands $\tt{steering} \in [-40, 40] $ degrees, where 0 is go straight, +20 turn 20 degrees left and -20 turn 20 degrees to the right.


The model is penalized when the car is driving in the areas too close to the left and right border of the lane. To reduce the penalty, steer fast out of the border area. Tune the model untill the penalty is below 25%. It would mean that the car drives less than 25% of the time too close to the borders. You gain more points if you can get the penalty smaller than 20% ? 

You may tune the parameters and run the model as many times a you like but
please be patient, the running of the simulation and the generation of the animation takes usually 20-60 seconds. 

```python
%matplotlib inline
%load_ext autoreload
%autoreload 2
```

```python
from IPython.display import HTML
from agentlibrary import autonomousCar as Car
```

Give your solution to the next cell. Remove the line saying `raise NotImplementedError()`

Current implementation of `steer()`-function returns steering angle 0, in all cases. Replace the function body with your own code. The function gets perceptions left and right. Left is true, if the car is too close to left border and right is `True` if the car is too close to the right border. Otherwise both signals are `False`.

Use the Python [if-elif-else](https://www.tutorialspoint.com/python/python_if_else.htm) structure and make function to return suitable steering angle in all conditions. Remember to indent the code logically, it is mandatory for Python to work.
 

<!-- grade_id: cell-38bf026327b2c0e9, points: 0, grade: false, solution: true -->
```python
def steer(left,right):
    return 0
```

<!-- grade_id: cell-eaa246f97c83c70f, points: 2, grade: true, solution: false -->
```python
# Testing of the operation of the car
# --------------------------------

# Instantiate a car simulation object
car=Car()

# Assign own function for steering the car
car.steer=steer

# Run the car simulation 
car.run()

# Make an animation of the simulation
HTML(car.anim.to_html5_video())

# !WARNING: the creation of the animation may take time. It should be finished 
#           in less than 60 seconds

```

<!-- grade_id: cell-7655b7495a1a21d8, points: 2, grade: true, solution: false -->
```python
""" Check if the penalty is less than 25% """
print("Penalty is %d %%" % car.getPenalty())
assert(car.getPenalty()<25)

```

<!-- grade_id: cell-676463bbc343bd5c, points: 2, grade: true, solution: false -->
```python

```

# Model based agent

The simple reflex agent shown in the previous example seems to work, but is far from optimal. Clearly we should try to turn the car parallel to the road, but we didn't have enough information for doing that. To improve the performance, we can add a state variable which tries to describe the orientation of the car. If we do not have any other sensors available, we cannot know the orientation exactly, but we may try to estimate it based on the feedback we get from the car simulator.

It seems that the car is initially approaching to the border of the road at the certain angle. Our steering actions converted it out from the border, but towards the opposite border with the steepness than in the beginning. If we needed to turn angle $\alpha$ to avoid the border, we need to turn angle $\alpha/2$ to the opposite direction afterwards, to align the car with the road. 

Try to use the orientation variable to perform some corrective steering actions after avoiding the first border. It is enough to just count how many steps you have steered to one direction, and after that steer equal amount of steps to the opposite direction, but half the angle used.

Now you should get penalty smaller than 10%.

Replace the function with your own implementation as you did above. In this time the code cell contains a global variable `orientation` which is accessed from the function. The reason for using a global variable is only that the value of the variable is kept between function calls. In this way you can easily count how many times you have already steered to the right, and then be able to compensate suitable amount after avoiding the first collision.

<!-- grade_id: cell-a5f394d0d32762de, points: 0, grade: false, solution: true -->
```python
orientation=0

def steer(left, right):
    global orientation
    return 0
# YOUR CODE HERE
```

<!-- grade_id: cell-ded028f87d50a015, points: 2, grade: true, solution: false -->
```python
# Testing of the operation of the car
# --------------------------------

# Instantiate a car simulation object
car=Car()

# Assign own function for steering the car
car.steer=steer

# Run the car simulation 
car.run()

# Make an animation of the simulation
HTML(car.anim.to_html5_video())

# !WARNING: the creation of the animation may take time. It should be finished 
#           in less than 60 seconds

```

<!-- grade_id: cell-c99f5b389541280a, points: 3, grade: true, solution: false -->
```python

print("Penalty is %d %%" % car.getPenalty())
assert(car.getPenalty()<10)

```

# Utility

Assume that you are in a house which has seven rooms. You enter in the lobby (room A) and there are two doors. One is for the kitchen (Room B) and the second one for the living room (room C). If you go into the kitchen, there are two other doors leading to the utility room (room D) or to the pantry (room E). In the living room there are also two doors, one leading to sleeping room (room F) and one to the wardrobe (room G).

You have 100 € of money and you know that a malicious robot will soon enter in the house from lobby (room A), and its goal is to find your money and take it. You know that the robot has so efficient sensors that it always finds the money if it is in the same room with them. Fortunately it does not search from all rooms. 

When the robot is in the lobby, it chooses between kitchen (B) and living room (C) at the probabilities shown in the graph below. Similarly, it selects between other possible rooms according to the probabilities shown deeper in the graph. The robot only goes from one room to the another according to the directions of the arrows.

Hide 100 € of your money in the rooms as you like, and then the simulation simulates the behaviour of the robot in ten different rounds, and calculates the average amount of funds that was not found in each round. Try to use the probabilities to hide your money in the rooms where the robot visit is less probable. Try also split your money to many different rooms, or try to put everything in just one room. What is the best strategy you can find? 

You can place your bet using for example the following code
`G.placeBet(['D', 'E'],[100, 0])`

The first argument for the placeBet function, `['D', 'E']` is the list of all rooms where you want to hide your money. The second argument `[90,10]` is the list of amounts of money you want to hide in the rooms you listed. In this example you would hide 90 € in room D and 10 € to to room E, and nothing to any other rooms. You can distribute the money in as many rooms any way you want, but the total sum must be always 100 €.

If you want to utilize more programming, you may generate random distributions of funds and run the simulation in the loop, until you find good enough solution (random search). In this case use `x=np.random.rand(4)`-function to generate four random numbers, then scale `x` so that the sum of the values is 100, and use the x-vector as a second argument to the function `G.placeBet(['D', 'E', 'F', 'G'], x)`. You may use for example `while`-loop, until the simulation result is good enough.

```python
from agentlibrary import UtilityGraph
G=UtilityGraph(7)
G.plot() 
```

<!-- grade_id: cell-7d854d9b05334da3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
G.placeBet(['D', 'E', 'G'],[0, 10, 90])
```

<!-- grade_id: cell-4088c12b4084c67e, points: 3, grade: true, solution: false -->
```python
moneyLeft=G.simulate()
print("Amount of money left = %d €" % moneyLeft)

if(moneyLeft<=75):
    print("This was not very optimal bet")
elif(moneyLeft>75) and (moneyLeft<85):
    print("This is good and acceptable. You could perform even better though")
elif(moneyLeft>=85):
    print("This was excellent bet")


```

```python

```